# ⌛️ Practice Final Exam

This Practice Final Exam is designed to test your mastery of the Python programming language. You will be asked to respond to a series of programming problems.


## Instructions

You are allowed to use anything with the python kernel to help you solve the problems. 

## Question 1:  Probability of Common Birthdays

Given a group of N people, what is the probability that at least one day of the year is the birthday of two or more members of the group? In this assignment, we will use simulations of groups of people with randomly assigned birthdays to find out.



The program template given below, my_birthday.py, defines the function do_birthday_simulation() which has three arguments. N is the number of people, M is the number of simulation trials to perform, and a_seed is the seed for the pseudorandom number generator. Note that the statement:

```python
birthdays = [random.randint(1,365) for j in range(N)]
```

assigns a random integer between 1 and 365 inclusive to each of N elements of a list called birthdays. Complete the function definition so that it returns True if at least one of those integers appears two or more times in the list.

The __main__ block is provided in the template so that you can run your program locally on your own computer for testing. It does not matter what the __main__ block contains when you submit; you will only be graded on the performance of your do_birthday_simulation() function.

```python
'''
D R E X E L   U N I V E R S I T Y
ENGR 131 -- Introductory Programming for Engineers

Programming Assignment 1

Written by Your Name Here
Term:  Winter 2020-2021
'''
import random
def do_birthday_simulation (N, M, a_seed):
    random.seed(a_seed)
    num_hits = 0
    ''' run M simulations and increment num_hits for each one that is a hit '''
    for i in range(M):
        ''' This statement generates the N random birthdays '''
        birthdays = [random.randint(1,365) for j in range(N)]
        ''' Your code goes here '''
    return num_hits / M

if __name__ == '__main__':
    N = 25
    M = 100
    a_seed = 1972227
    print('For {:d} simulations of groups of {:d} people with seed {:d},\nthe probability of a common birthday is {:.2f}'.format(N,M,a_seed,do_birthday_simulation(N,M,a_seed)))
```

## Question 2: Manipulating DNA

Part A:
Given: A DNA string, s, where the alphabet of DNA is from the set {'A','C','G','T'}

Return: Four integers (separated by spaces) counting the respective number of times that the symbols 'A', 'C', 'G', and 'T' occur in s. On the next line, a float with the CG% (which is the combined percentages of Cs and Gs in the string). The percentage is from 0 to 100 and NOT a fractional percentage.

Sample Input: ATTTTAAGGAGTTTAAAATGGATAAGAAGCAAGTAACGGATTTAAGGTCGGAACTACTCGATTCACGTTT

Sample Output: 25 8 16 21

34.3

Part B:
Given two strings s and t of equal length, the Hamming distance between s and t, denoted
, is the number of corresponding symbols that differ in s and t.

Given: Two DNA strings s and t of equal length .

Return: The Hamming distance $d_H(s,t)$.

1st Sample Input:

GATATCGTCTGGGACCT

CATCGCATTTACGGCCT

Sample Output:

8



2nd Sample Input:

GATATCGTCTGGGACCT

CCATCGCATTTACGGCCT

Sample Output:

Error: Sequences Length Mismatch



You need to write a function for each part.
The declarations for both functions are given in the template.

You may use more functions if you find that helpful for your implementation, but only these two will be explicitly checked.



```python
slist=input()
seq=list(slist)

def countBases(seq):
    # return count of As, count of Cs, count of Gs, count of Ts, and Percentage of Cs and Gs combined
    
    return '{:d} {:d} {:d} {:d}\n{:2.1f}'.format(A, C, G, T, cg_percentage)
    #return A, C, T, G, cg_percentage

def hammingdistance(seq_a, seq_b):
    #return the number of differences between the two sequences
    #If the sequences are not the same length, return "Error: Sequences Length Mismatch"
                
    return hamming_distance

if __name__=="__main__":
    # optional - add any test code here inside this block. we do this so that when the zyBooks tests includes 
    # your files to call your functions directly that the testing code in this block is not invoked
    
    #basecount=countBases(<input sequence here>)
   
    pass
```

## Password modifier

Many user-created passwords are simple and easy to guess. Write a program that takes a simple password and makes it stronger by replacing characters using the key below, and by appending "!" to the end of the input string.

i becomes 1
a becomes @
m becomes M
B becomes 8
s becomes $

Ex: If the input is:<br>
```mypassword```<br>
the output is:<br>
```Myp@$$word!```<br>




```python
word = input()
password = ''

''' Type your code here. '''
```

## Parsing dates




Write a program to read dates from input, one date per line. Each date's format must be as follows: March 1, 1990. Any date not following that format is incorrect and should be ignored. The input ends with -1 on a line alone. Output each correct date as: 3/1/1990.

Hint: Use string[start:end] to get a substring when parsing the string and extracting the date. Use the split() method to break the input into tokens.

Ex: If the input is:
```March 1, 1990```<br>
```April 2 1995```<br>
```7/15/20```<br>
```December 13, 2003```<br>
```-1```<br>
then the output is:<br>
```3/1/1990```<br>
```12/13/2003```<br>

```python
def get_month_as_int(monthString):

    if monthString == 'January':
        month_int = 1
    elif monthString == 'February':
        month_int = 2
    elif monthString == 'March':
        month_int = 3
    elif monthString == 'April':
        month_int = 4
    elif monthString == 'May':
        month_int = 5
    elif monthString == 'June':
        month_int = 6
    elif monthString == 'July':
        month_int = 7
    elif monthString == 'August':
        month_int = 8
    elif monthString == 'September':
        month_int = 9
    elif monthString == 'October':
        month_int = 10
    elif monthString == 'November':
        month_int = 11
    elif monthString == 'December':
        month_int = 12
    else:
        month_int = 0

    return month_int


user_string = input()

# TODO: Read dates from input, parse the dates to find the one
#       in the correct format, and output in m/d/yyyy format

```

## For loops over parallel containers with zip()
In chemistry, equilibrium constants can be expressed as products of species concentrations each raised to a power corresponding to that species' stoichiometric coefficient in chemical reaction. For example, for the reaction

$$ A + 3B \rightleftharpoons 2C $$

the equilibrium constant is

$$ K = \frac{[C]^2}{[A][B]^3} $$

where [ ] indicate species concentration.  Note that stoichiometric coefficients of reactants (species on the left-hand side of the arrows) are negative, and stoichiometric coefficients of products (species on the right-hand side of the arrows) are positive.  That is, we could also write the equilibrium constant for this reaction as

$$ K = [A]^{-1}[B]^{-3}[C]^2 $$

Generally, for a set of species 

$$ I \in (A, B, C, \dots ) $$

and their corresponding stoichiometric coefficents:

$$ \nu _i \in ( \nu _A, \nu _B, \nu _C, \dots ) $$

we can express the equilibrium constant as a **product** expansion over N species in the reaction:

$$ K = \prod _{i=1}^{N} [I]^{\nu _i} = [A]^{\nu _A} [B]^{\nu _B} [C]^{\nu _C}\dots $$

In this lab, you will write a program to compute the value of K given two lists:  the first is a list of values of concentrations of species, and the second is the corresponding list of stoichiometric coefficients.  Each list is input as a comma-separated string, and the two lists are input on separate lines.  For example, if the input is

```
0.5,0.75,1.3,0.8
-1,-2,2,1
```

the program will compute (0.5)<sup>-1</sup>(0.75)<sup>-2</sup>(1.3)<sup>2</sup>(0.8)<sup>1</sup>, generating the single output

```
4.80711
```

A key aspect of this lab is that the program should be able to handle **any** number of species; this is only determined by the number of comma-separated elements in the input list.

Python provides the `zip()` built-in class to permit iteration over multiple containers in parallel.  When provided multiple containers, `zip()` yields tuples of container elements one at a time, like this:

```
>>> a=[1,2,3]
>>> b=['c','d','e']
>>> zip(a,b)
<zip object at 0x00000245990F84C0>
>>> list(zip(a,b))
[(1, 'c'), (2, 'd'), (3, 'e')]
```

Because a `zip` is iterable, you need not cast it to a `list` to iterate over it:
```
>>> for i,j in zip(a,b):
...     print(i,j)
... 
1 c 
2 d 
3 e 
```

Use a `for` loop over a `zip` of the two input lists to compute a value for K.



```python
C = list(map(float,input().split(',')))
nu = list(map(float,input().split(',')))

# Your code goes here; assign a value to K, the equilibrium constant

print(f'{K:.5f}')
```

Tests for q5 - Making a Graph (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [3, 3, 3, 6, 14, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q5_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_1'
>>> max_score = 3
>>> score = 0
>>> if np.__version__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': np.__version__}
>>> responder.add_response(response)
>>> assert np.__version__ is not None, 'numpy is incorrectly imported'
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_2'
>>> max_score = 3
>>> score = 0
>>> if plt.__package__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': plt.__package__}
>>> responder.add_response(response)
>>> assert plt.__package__ is not None, 'plt is incorrectly imported'
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_3'
>>> max_score = 3
>>> score = 0
>>> if callable(draw_card_eq):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': callable(draw_card_eq)}
>>> responder.add_response(response)
>>> assert callable(draw_card_eq), 'draw_card_eq is not a function'
```
failure_message: draw_card_eq is not a function
success_message: draw_card_eq is callable.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import matplotlib.pyplot as plts
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_4'
>>> max_score = 6
>>> score = 0
>>> if isinstance(plot_[0], plts.Line2D):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': isinstance(plot_[0], plts.Line2D)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> scorer = submit_score()
>>> assert isinstance(plot_[0], plts.Line2D), 'You did not return a plot object.'
```
failure_message: You did not return a plot object.
success_message: You successfully returned a plot object.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_5'
>>> max_score = 14
>>> score = 0
>>> if np.isclose(sum(plot_[0].get_xdata()[::3]), 1.6561871562318847e-07):
...     score += int(max_score / 2)
>>> if np.isclose(sum(plot_[0].get_ydata()[::4]), 1.2500506348082654):
...     score += int(max_score / 2)
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': sum(plot_[0].get_xdata()[::3])}
>>> responder.add_response(response)
>>> assert np.isclose(sum(plot_[0].get_xdata()[::3]), 1.6561871562318847e-07), 'The values for x_data in the plot are incorrect. '
>>> assert np.isclose(sum(plot_[0].get_ydata()[::4]), 1.2500506348082654), 'The values for y_data in the plot are incorrect. '
```
failure_message: The values for your plot are incorrect.
success_message: The values for your plot are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_6'
>>> max_score = 6
>>> score = 0
>>> if plot_[0].get_color() == 'red' or plot_[0].get_color() == (1.0, 0.0, 0.0, 1.0) or plot_[0].get_color() == 'r':
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': plot_[0].get_color()}
>>> responder.add_response(response)
>>> assert plot_[0].get_color() == 'red' or plot_[0].get_color() == (1.0, 0.0, 0.0, 1.0) or plot_[0].get_color() == 'r', 'The line on the plot is not the correct color'
```
failure_message: The plot is not red.
success_message: The plot is red as expected.


Tests for q6 - Exploring DataTypes (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [3, 3, 3, 3, 6, 6, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q6_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_1'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDQ71dNmKvaPlZOj3cLT2FVfh66b09i0yV2Fl08bcFY0K_Io7ju7utXkX3cCmsb9HMsHTAHERX0X48acRp1MGHYXxSg==')
>>> if str(int_) == out and isinstance(int_, int):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': int_}
>>> responder.add_response(response)
>>> assert isinstance(int_, int)
>>> assert str(int_) == out, 'int_ incorrectly implemented.'
```
failure_message: int_ incorrectly implemented.
success_message: int_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_2'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRABl3uUfW0G0yLb5kCpVpz5XBRYUjowgJ9_c9CNAGielysMhb5cp72zDF0Gi9abZrATkUoZfUGv7UJp65Cm7WvjiQ==')
>>> if str(float_) == out and isinstance(float_, float):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': float_}
>>> responder.add_response(response)
>>> assert isinstance(float_, float)
>>> assert str(float_) == out, 'float_ incorrectly implemented.'
```
failure_message: float_ incorrectly implemented.
success_message: float_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_3'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRDR42csyMOkhel0VkaN7JXse5_3OFTam1ryIyiJvEdIcfDCfocojIiGZU2PPJQm-rrfBNXLyzMQ-QJgmZi7qEdthA==')
>>> if str(bool_) == out and isinstance(bool_, bool):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': bool_}
>>> responder.add_response(response)
>>> assert isinstance(bool_, bool)
>>> assert str(bool_) == out, 'bool_ incorrectly implemented.'
```
failure_message: bool_ incorrectly implemented.
success_message: bool_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_4'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRECpI3kPdhirMzwLAddFoJKDXs7mXJ5AWPYJTx_gTKGYY9VR0REjaPKLCAc0MQ5UuSr4cMcBBc8zSVg2GT6LiDvLg==')
>>> if str(string_) == out and isinstance(string_, str):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(string_)}
>>> responder.add_response(response)
>>> assert isinstance(string_, str)
>>> assert str(string_) == out, 'string_ incorrectly implemented.'
```
failure_message: string_ incorrectly implemented.
success_message: string_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_5'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRFLoBpFP_06PNjYcIdyOpLGFYYboKIjMjTAu936eLjoKByeaF5XTyYcvGh5PWyVVU1kFUqWj4YeqbRDfMTRpPYXGw==')
>>> if str(list_) == out and isinstance(list_, list):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(list_)}
>>> responder.add_response(response)
>>> assert isinstance(list_, list)
>>> assert str(list_) == out, 'list_ incorrectly implemented.'
```
failure_message: list_ incorrectly implemented.
success_message: list_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_6'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRGEX2Cl-qHsEg1NFnwyZSHVLsjlc1JJWn4Lj76LT7dsEoNqlBpXQry0zKyFS2oCSzdJbYSm95SKxB96pM_HHsyuNw==')
>>> if str(tuple_) == out and isinstance(tuple_, tuple):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(tuple_)}
>>> responder.add_response(response)
>>> assert isinstance(tuple_, tuple)
>>> assert str(tuple_) == out, 'tuple_ incorrectly implemented.'
```
failure_message: tuple_ incorrectly implemented.
success_message: tuple_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_7'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRG1TzclExNe_t7qeJgaHYC9uWufJOfW5u7L198U2jtPxlFPQTnL9qj0MFHrNC7pIUfpz7od9FbuyQy6nuNnEFvlk2ZWGw3xRiU_wh4fVtRuV-Y=')
>>> if str(dict_) == out and isinstance(dict_, dict):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(dict_)}
>>> responder.add_response(response)
>>> assert isinstance(dict_, dict)
>>> assert str(dict_) == out, 'dict_ incorrectly implemented.'
```
failure_message: dict_ incorrectly implemented.
success_message: dict_ correctly implemented.


Tests for q7-Building a Thermometer (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [6, 6, 8, 8, 8]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q7_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_1'
>>> max_score = 6
>>> score = 0
>>> if isinstance(Thermometer, type):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(isinstance(Thermometer, type))}
>>> responder.add_response(response)
>>> assert isinstance(Thermometer, type), 'The Class Thermometer is not implemented.'
```
failure_message: The Class Thermometer is not implemented.
success_message: The Class Thermometer is implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_2'
>>> max_score = 6
>>> score = 0
>>> t = Thermometer(25)
>>> if hasattr(t, 'temperature_c') and t.temperature_c == 25:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(t.temperature_c)}
>>> responder.add_response(response)
>>> assert hasattr(t, 'temperature_c'), "Thermometer does not have 'temperature' attribute"
>>> assert t.temperature_c == 25, 'Thermometer initialization did not set attribute temperature correctly'
```
failure_message: The Class Thermometer initialization is implemented incorrectly.
success_message: The Class Thermometer initialization is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_3'
>>> max_score = 8
>>> score = 0
>>> t = Thermometer(25)
>>> t.set_temperature(50)
>>> if t.temperature_c == 50.0:
...     score += 4
>>> t_ = Thermometer(25)
>>> t_.set_temperature('abc')
>>> out = str(t_.temperature_c)
>>> if out == 'fail':
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([t.temperature_c, t_.temperature_c])}
>>> responder.add_response(response)
>>> t = Thermometer(25)
>>> t.set_temperature(50)
>>> assert t.temperature_c == 50.0, 'set temperature does not function properly'
>>> t = Thermometer(25)
>>> t.set_temperature('abc')
>>> out = str(t.temperature_c)
>>> assert out == 'fail', 'set temperature does not return the correct response when a unsuitable input is provided'
```
failure_message: set_temperature can not correctly handle non numeric inputs.
success_message: set_temperature can correctly handle non numeric inputs.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_4'
>>> max_score = 8
>>> score = 0
>>> t = Thermometer(0)
>>> if t.to_fahrenheit() == 32.0:
...     score += 4
>>> hold = t.to_fahrenheit()
>>> t.set_temperature(20)
>>> if t.to_fahrenheit() == 68.0:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([hold, t.to_fahrenheit()])}
>>> responder.add_response(response)
>>> t = Thermometer(0)
>>> assert t.to_fahrenheit() == 32.0, 'Incorrect conversion to Fahrenheit.'
>>> t.set_temperature(20)
>>> assert t.to_fahrenheit() == 68.0, 'Incorrect conversion to Fahrenheit.'
```
failure_message: Incorrect conversion to Fahrenheit.
success_message: Correct conversion to Fahrenheit.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_5'
>>> max_score = 8
>>> score = 0
>>> conditions = [hasattr(water_sensor, 'temperature_c'), hasattr(water_sensor, 'to_fahrenheit'), temp_c == 12, temp_f == 53.6]
>>> statements = ["Thermometer does not have 'temperature' attribute", "Thermometer does not have 'to_fahrenheit' attribute", 'Incorrect temperature for temp_c.', 'Incorrect temperature for temp_f.']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Incorrect conversion to Fahrenheit.
success_message: Correct conversion to Fahrenheit.


Tests for q8-Determining a Price for a Circuit (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [0, 1, 0, 6, 15, 12, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q8_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_1'
>>> max_score = 0
>>> score = 0
>>> conditions = [parts_cost['10-ohm resistor'] == 0.025, parts_cost['100-ohm resistor'] == 0.03, parts_cost['1-pF capacitor'] == 0.05, parts_cost['10-pF capacitor'] == 0.12]
>>> statements = ['10-ohm resistor not correctly defined', '100-ohm resistor not correctly defined', '1-pF capacitor not correctly defined', '10-pF capacitor not correctly defined']
>>> for condition in conditions:
...     if condition:
...         score += 0
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: parts_cost is incorrect. Please make sure you did not modify the dictionary.
success_message: parts_cost is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_2'
>>> max_score = 1
>>> score = 0
>>> if np.__version__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(np.__version__)}
>>> responder.add_response(response)
>>> assert np.__version__ is not None, 'numpy is incorrectly imported'
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_3'
>>> max_score = 0
>>> score = 0
>>> conditions = [circuit_parts['10-ohm resistor']['quantity'] == 10, circuit_parts['10-pF capacitor']['quantity'] == 3, circuit_parts['1-pF capacitor']['quantity'] == 7, circuit_parts['Custom Processor']['quantity'] == 2, circuit_parts['Custom Processor']['price'] == 12]
>>> statements = ['10-ohm resistor quantity not correctly defined', '10-pF capacitor quantity not correctly defined', '1-pF capacitor quantity not correctly defined', 'Custom Processor quantity not correctly defined', 'Custom Processor price not correctly defined']
>>> for condition in conditions:
...     if condition:
...         score += 0
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: circuit_parts is incorrect. Please make sure you did not modify the dictionary.
success_message: circuit_parts is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_4'
>>> max_score = 6
>>> score = 0
>>> conditions = [price_calculator.__name__ == 'price_calculator', price_calculator.__code__.co_argcount == 3]
>>> statements = ['you did not name price_calculator correctly', 'price calculator does not have the correct number of inputs']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([price_calculator.__name__, price_calculator.__code__.co_argcount])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: price_calculator call is incorrect.
success_message: price_calculator call is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_5'
>>> max_score = 15
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkDz6Y_aio_akqMGkihfI5zKR2KkUg0yXVVHNrVV82qunlAwRjVZN2R2OYzUqkBBLBG-qd4t_2Vpx-iHxYPLMGTYQsFA==')
...     if str(price_calculator(parts_cost, circuit_parts, 0.2)) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(out)}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert str(price_calculator(parts_cost, circuit_parts, 0.2)) == out, 'price_calculator  computation is incorrect.'
```
failure_message: price_calculator computation is incorrect.
success_message: price_calculator computation is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_6'
>>> max_score = 12
>>> score = 0
>>> parts_2 = parts_cost.copy()
>>> parts_2['10-ohm resistor'] = 10
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD3dtfMyEiX9AYR1pum4OZ5xSbbZ367hae3wxg5SSAcor4m5CXO6u1OthJkNgyGQUotRmCY2v18f4rjxK7Z3kbCNC3A==')
...     if out == str(price_calculator(parts_2, circuit_parts, 0.2)):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(out)}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert out == str(price_calculator(parts_2, circuit_parts, 0.2)), 'Your code does not use the parts_cost correctly'
```
failure_message: Your code does not use the parts_cost correctly.
success_message: Your code correctly uses the parts_cost.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_7'
>>> max_score = 6
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD3D3Z-O5XJGgN67FRSetS5oiOetUTzWxQFbw_OTKnrjBEhnelTKcH0d22z4AneFL8pYrtOgSg0Q2qYc0A4dI-VdzzcV6pXrv5aKw-FAVmhHMZ7U=')
...     if print_total(price_calculator(parts_cost, circuit_parts, profit)) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(out)}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert print_total(price_calculator(parts_cost, circuit_parts, profit)) == out, 'The string being printed is not correct'
```
failure_message: print message is incorrect.
success_message: print message is correct.


Tests for q9-Plotting and Fitting (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [3, 9, 4, 6, 3, 9, 3, 4, 4, 3, 3, 6, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q9_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_1'
>>> max_score = 3
>>> score = 0
>>> conditions = [np.__name__ == 'numpy', plt.__name__ == 'matplotlib.pyplot', 'curve_fit' in dir()]
>>> statements = ['numpy not imported correctly', 'matplotlib.pyplot not imported correctly', 'scipy.optimize.curve_fit not imported correctly']
>>> for condition in conditions:
...     if condition:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(plt.__name__)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Your input statements are incorrect.
success_message: Your input statements are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_2'
>>> max_score = 9
>>> score = 0
>>> x_test = np.array([0, 1, 2])
>>> (a_test, b_test, c_test) = (1, 1, 1)
>>> expected_output = np.array([0.60653066, 1.0, 0.60653066])
>>> conditions = [np.allclose(gaussian(x_test, a_test, b_test, c_test), expected_output)]
>>> statements = ['gaussian function is not implemented correctly']
>>> for condition in conditions:
...     if condition:
...         score += 9
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([*gaussian(x_test, a_test, b_test, c_test)])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: gaussian function is not implemented correctly
success_message: gaussian function is implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_3'
>>> max_score = 4
>>> score = 0
>>> conditions = [np.allclose(x_data[0], -20), np.allclose(x_data[-1], 20), len(x_data) == 100, np.isclose(sum(x_data[3::3]), 19.99999999999997)]
>>> statements = ['x_data is not defined correctly', 'x_data is not defined correctly', 'x_data length is not correct', 'x_data is not defined correctly']
>>> for condition in conditions:
...     if condition:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: x_data is not defined correctly.
success_message: x_data is defined correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_4'
>>> max_score = 6
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-NzYrMglFaqNyS-zlkoQo-gwxc1YCD39uP-uNVNborNwKBB7YKbPEpayy9rq7hUGxJhZWrWmXBDQ14fxazkpMi5v-8ArClYcJqu2ZJvUs2LKzPwphZZ63pfZBxs8IK0DIw-EnU5mRoiaMT3YoTAD7rSGUrnDdqM44YGV3b_Gyt7mYE0n_obgjbXngI0DbuJzExiXeA4KQNpSsKJD6x870owlz7wNXADMhSrU03suFLlUx1pgGwJsg-7lQhUocs64ZahOJvSfABgqu_Ml_SDkaudvTyHTFCPy-bnq3xn-Cc4d9jmsdOxUmqMvl3xLZw1ALt6YrMBdlNZh0hBgjrT4iVP9tySJBG8dnov7TLz9_PuWRGRulKG-tYZU9XJ8uuFC09mXpDLE49gdlCDp_jYbvYfWX6rPlgBd7pXyQ2YWAKajlRzuJ1Q7TrXSVm_Xbj2eF48_iehGgz2cuzg0YARRKXvrg3JFdRn-jRhN9Pop4SoJVEcIbh6p6AmU0u2rUTFDN5UazGdLSWav8g2xQyHCm8_2i9il6lVF3_H4E5gdycNi6eb4hUV0i7XY1rUzfulJYmJzWzQk3KgbR54In9AfLWvgHY0IcBduStXBmpMfwIG7NMbbSPhElQ0D6gXjA3F7zpkahXZ3S0ihF0bfHvY32pIfRVORtMT811bYif8OPTJH2i3YcHiaS7ElIUJXkb4dnJTOOVCIB08EV0ONZgZHwtpsFCCwCFvRTouorcfz1SXs9a0vtjtau3CgBI3cl0yAGGuiazfipZAEsbB68QedFffR_QRl0ZBN_IfGIMC8ksxUw3R9J46tUET_doctLBvd7LFbDTfkECNoKCh-E5x-POLyjeWgscHe64WSRNkYgKaKTWP70Y9OdYg2CpCZseE55ovOqgNjTVCMlRhuQWO9zuBSLAU3NOfL-ms4n9AWrx75Wsg9gxkIGVm36HZ-yKO_wJDcrRRQqtYaS0noeqEdo3-Gmj4ldMgPOATkNXKLoFJ7P6FfnNYUiY7mPykn1_1tqlBEhVvSXrAteRNr6E_8s6lgFr4l2kDU8s8_XSgKu_wazGqw45WVMiO5Zgsh2SXh4VnnUE0LFubdThZsKCplGr_uJK2IFc0jiWlSid8rE2GiLU69Y0mVkwod2PQOfIfJRzeNkAA9raRbs5CkFO8XFVgLGjKWGIJw_8zKncKHZ-LgK0PVFJ2zMBMuN8fXofonS3v7Zm85gj16MGtfwIOVhvODG4BrOkZvI43TCM5pPRq1-5nqyJfJ6gpC0SWWwMrEF7YSZ64YIA5YmVlaquaB5Fam4p5O3EN8m83aQ4AcCBcRZAaimCtPIrIhK7SS394wkfQLFb8q_WUJe5RBHm3PoVsbjac3Vlu7rQTVByxSf9KdFNdtKTXJsMzLL0743ZMBmm4RA5UlQWE15AFeIyiSXP8B0TgFKe9J5g8W-eT8KWyqRw-x00_7lF0e-4tcB1B6CLOuk4kpkZk0WkucN0S1IzZRZsjXcOwmsoNDqFdpx7NrO08ZgmXqkSqwLCo99q0C6gpk2oX3-BBw9cDOgvFIpYE4A==')
...     if str(y_true) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(y_true)}
>>> responder.add_response(response)
>>> assert str(y_true) == out, 'y_true is incorrect'
```
failure_message: y_true is incorrect.
success_message: y_true is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_5'
>>> max_score = 3
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-UVvxc6wECqC_U9taa6Art0EtaZg2i_WB0sp4FHCf3Ggyed8uUrljAkCPxg-AwYxN2jnqaPVgMh57F1OX1ftyjzR1SfhzBk5UAkxPaLIFHteLZY_7oK67GesWvX-jys-fhfKLktB_4ba54nXaVQP3K2xb60izc-42Y4VvddXu9Iggo5dy7VcJP8-6Ix_un8KqJo2J4CKYXFni2El-O9zrvvc1EcNYgvGdpzxFFeUHEHAP0AEqTp-Ow5D7T0x2iGcEPXtDUcmbgMG_t-HAQRr8ky9DNyuKx9k13pWT85O3jpLAtExhBYk2yXtzS9iU4pQEYGpYhdB2HQosUQcV5rXdgsC6pGcWlWhTab2uG32_Lt8-Y0jlb4VgALBAtxJGXuevGpJ7Pt1fzt7pjHikCVHcMbmkhLaWGzZq1uf6d6mbA-lxo1wGif7-ryFB4TTyze2HoxcGGgLI7vtuHgduBSd397OTTuS9PMMnR7oIXBZprj_NJrSmI79N4mMDlkWz-baZikhLw6LbThxinvd4HX6aQPVm3bPlV9M5Yc3lSqr2YQimUKqONnZGKxzcQx4TxfeRjdv8mS49Qey1KRU5zUlBS11NX-ScrqIFDAwFBsq6OpMSqd80rZquYE5E-IsHr0g4McffuDal20JZqXtP_kYMaT3Iv_KCa__9TKmpMHKJpQ44qhvzuBkqA59npHU05s5gbr5OpPwug1W_H7WTgjRYD-IaDU9cxxKWs9erfEkcCMh7peMvnBVwu2IY_W1u7ocNxXiMR2BJX6FbGO2P12t7D8vz3lyCJFHhkYggO06v9UF0L3-eX186gXNdegYnML3kVevZbOkRHFxvU2hoZ6dzLW8Z8SgaO9H1-Pac5Ima1Kh36EF1GH-93QGnSf6w2feXG9yByfSSbRharfHT6-79nYgl3EgGL0wo8imUaN8uwRhkxma8myGCAsNbkPKPiltTXfyv8yKEjc-qC5UivinYtFYYkhiDqIQvQDB4YURk0p4MxesM1rOgq6aML_jc7haRUcUes4ynGERvVt_Nc3bV03--BIzTZx4D2q3OMEUt25CyTU9oUr6r4C_JsrxhtK2d5U9qrLY6b-W-F-lBI3ZrRh_HcwHR9Hx7YTA7vRcXaOyYB-fUWjwXa2Gbxj1s9iEpzWf8lEisQk40FUhc2jSmtiZEn6MO3dO_YxJ3s9TE1Tmq9aK7gmb3zMbZn2gQyuHxzuvnUeJBi-DelNLvQXP7C3GYyMjA6TuajEzNY6VSVE7sKgQyw_vjE0QteWtkUpODlAbpT42M6laN3ahNeT7l2WE0D-wvjEHzflGOF7udat6QccodlHUge4DAJSS4vaTWs5J-aoSCmPNrYN65W3CIokMhxXS0VZ0n26ObdU384YVKdOTtUn796D7Mdk0ll0ImLLs_1RcihuZPkmEhE2rnWUINRBOLAyoQz6RrNGe-tC-uA3XNsxYbDQ9ReBPlU_R_BbkkgEzKtMV7juGtddtKDF-1lzhoahFX33CNMdWcuq189G87JsEED_tdmSloByFx3sEExB0uuj0--pnbH-CV1674xJ8x2WarCevNkOts1Aa0tDyE9v9v8NmX6cIKHsT8rkIhStCE1_Xknn2uMGlqNPp_5k30P5_mKeth6tti6uH-MMNO_qQ8V82Ki7X5u3HOmPgAa3_uJGD-JVvREQ8uNw8SVFpCmO-X207mkqYEFXO_Q=')
...     if str(y_noisy) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(len(y_noisy))}
>>> responder.add_response(response)
>>> assert str(y_noisy) == out, 'y_noisy is incorrect'
```
failure_message: y_noisy is not implmemented correctly
success_message: y_noisy is implmemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_6'
>>> max_score = 9
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-VjEnvkOx1q5hIc6OiMyPE7iayGWAtZR_Kry0mvrQbCr0ffV9Lrck5j673YFnlWxRjo9qR44gKpYzUOzGhmCBpXKZ2Dec1hHWXhY4szHZZuirxAWmfMc66dj5dpA1bzjYFG')
...     chars_to_remove = ['[', ']', '\n']
...     for char in chars_to_remove:
...         out = out.replace(char, '')
...     out = [float(x) for x in out.split(' ')]
...     if np.allclose(popt, out):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([*popt])}
>>> responder.add_response(response)
>>> assert np.allclose(popt, out), 'popt is not returning the correct values'
```
failure_message: popt is not returning the correct values
success_message: popt is returning the correct values
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_7'
>>> max_score = 3
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-WzDT72FrJF_4PrNsTIhnppKucv8pAjRnotXNrslY7YlSdDyqgiXabo32_bwcFXYkokXar7Vl-Bl1cN4lD70-PDCbfDPsxVDcGs25T4Z8IV-i6P7SAwaZ-k00NmVsJ4sF62ClW9ssomSd0jgYtbgy4J5sF9_KO-6s0ke6_Pfi1mrx8c90J5ZYsAwB9NKm45dZ-IWfjH955EkoxoPQcdqBQ16RG1J0k5DqGaP6hpRpR4mRANk50AFcgkX9PLqnKr3Z3AH1s_ugMEvza1GRlj1aVmaeF0JcHsVw1MxpZLLhPK0ikm8TcyX7cy6O8wi3xvjk5By8SFT96a0pqK5ObXGIuFAN6XlBH56yic8ttcRBODzaRiH_97wGvi8eUP0NNGOlMZFb_mPc7E9URWZCnJ0Krr57iOj5ALWDPqtMysnRGg4VCI9l6ZAFv0dP_HqbZT7U7TK_dFVAWNWv5MPeyg8rcGInJ2ygH1o1G8jRJRSbEcHQj3-ZU2K7DCFqc6xg2EO_31Dz6VLWnEBLVW2OUIth16vPCcjbmJgGuW0mgRulbogMNLFEXpGKJwi8LimBL_gZ0pee4vd8nZDKaRR8pSoLXsHyTHf-O8vXHzji4RiUF7VSReSrbpai1xoC7uyztjO1Vjce1bMf1pP7e3JNIzpJ7sfAlP0PPiERdG58Xe_2Qx_oWFMtQeDt-ViS4AWK76KlWgsiX0cFqWs3Dj_KMs1hmtywbEZtLkUjEUH6rclgw482Tp44fFQcUqBOOD92DJkGjmEr-FcVJtlu8R71hafSMBnOJ_cEUvZL-P5ljNUoUoM1I7REUK9hs-hiSYEyXbyCfv5dRSoXHOdNv0nDKoEcVWPYZPP3E74Qafmm47NHUNbcG3pMGXXMsOea7l5SnL_gfnYjNcvUIDEzZDXen2moiSP0P5EIeFTSItTdjSnM-TAMmDZlFSF3XtXD8vagpPb_-QdfPNOkkx1MQNjVXcKqDVPUwyuDkrqm0QFp2XrRW8Lwb7v5mENaD3HfZMlnpuBYAv9sXDtkb7Km30U9NgYExvScKEnBfbd61peROkRYXJiTA8uvx3xgWAeu47lCYNJ0jKomOQlJOVXH86TnPEjMKDQ9UUcFTJGcuQLaNfSPdSPqkf1gabg3e6VW4b--PSkhpss9OboUp1UOEo0KFnlj4Tl-zu8rwuWEUwlBUDYIyuKwWBYD_DpC62kBfCcuMuaM1xvbbMoedQLTm0OA0PIR_dP5nrW4IkGtnpbz43rqBMFCcGZ2un8KbCRT1mh8dB5nNdrtxl3dtBEB_nsl-M8gIqCruIfNHkoPx77SkXHT3zNnI66oYqgBpMCrTjShWB-xNtlRrKHqL0q9RpgPF9SJphcLBQ07eqV4s7TrJtfrhw41WuzLaCR7KYb_vsMj78KNrZmuaTDgGGu4lediwCYGOFA3I-jIxSj2jwYfDcrGKW9QUbDK5RTZ0_rEt6SwA5r7NpadXtv-RsGBycn-_vTNPQ_5PGIYGyY3GkQgVjiq0nBX1xI39PqmMKc0eyGKAH6pu2uhCw_FAe3GSSHzMke-zzIA==')
...     chars_to_remove = ['[', ']', ' \n']
...     for char in chars_to_remove:
...         out = out.replace(char, '')
...     out = out.replace('  ', ' ')
...     out = [float(x) for x in out.split(' ')]
...     if np.allclose(y_fit, out):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([*popt])}
>>> responder.add_response(response)
>>> assert np.allclose(y_fit, out), 'y_fit is returning incorrect values'
```
failure_message: y_fit is returning incorrect values
success_message: y_fit is returning correct values
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_8'
>>> max_score = 4
>>> score = 0
>>> conditions = [plot_1.axes.get_xlabel() == 'x', plot_1.axes.get_ylabel() == 'y']
>>> statements = ['xlabel is incorrect', 'ylabel is incorrect']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_1.axes.get_xlabel(), plot_1.axes.get_ylabel()])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: The plot labels are incorrect.
success_message: The plot labels are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_9'
>>> max_score = 4
>>> score = 0
>>> conditions = [plot_1.get_label() == 'Noisy Data', plot_2[0].get_label() == 'Fitted Gaussian']
>>> statements = ['Noisy Data', 'plot_1 is not labeled correctly', 'Fitted Gaussian', 'plot_2 is not labeled correctly']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_1.get_label(), plot_2[0].get_label()])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: the plots are not labeled correctly
success_message: the plots are labeled correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_10'
>>> max_score = 3
>>> score = 0
>>> conditions = [plot_2[0].get_color() == (0.0, 0.5, 0.0, 1) or plot_2[0].get_color() == 'g' or plot_2[0].get_color() == 'green']
>>> statements = ['gaussian fit is not the correct color']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_2[0].get_color()])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: gaussian fit is not the correct color
success_message: gaussian fit is the correct color
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_11'
>>> max_score = 3
>>> score = 0
>>> conditions = [plot_1.axes.get_legend() is not None]
>>> statements = ['The legend is not correct on the plot']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_1.axes.get_legend() is not None])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: The legend is not correct on the plot
success_message: The legend is correct on the plot
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_12'
>>> max_score = 6
>>> score = 0
>>> conditions = [all(plot_2[0].get_xdata() == x_data), all(plot_2[0].get_ydata() == y_fit)]
>>> statements = ['the x data for plot_2 is not correct', 'the y data for plot_2 is not correct']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: plot_2 does not have the correct values.
success_message: plot_2 is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_13'
>>> max_score = 6
>>> score = 0
>>> conditions = [np.array_equal(plot_1.get_offsets()[:, 1], y_noisy), np.array_equal(plot_1.get_offsets()[:, 0], x_data)]
>>> statements = ['the y data for plot_1 is not correct', 'the x data for plot_1 is not correct']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: plot_1 does not have the correct values.
success_message: plot_1 is correct.
